# SLU 17_1 - Exam Prep I

## Batch 1 Python Exam 

This is the Python admission exam from the Lisbon Data Science Starters Academy - Batch 1. The allocated time for this exam was 2 hours. If you plan to take the admission exam this year, it's a good idea to measure the time you need to solve this notebook. You can see how far you get in 2 hours and also how much time you need to complete the notebook. At the same time, it's a good idea to not rush through the exercises and rather concentrate on working thoroughly.
At the end, you submit the notebook to the portal as usual. You can submit as many times as you like.

```python
#used for evaluation
import inspect
import math
import re
import json
import sys
import hashlib

def _hash(s):
    return hashlib.blake2b(
        bytes(str(s), encoding='utf8'),
        digest_size=5
    ).hexdigest()
```

## Question 1

Complete the function `import_data`. This function should:
- Open the **countries.json** file that is located within the folder name **data** ;
- Use Python’s module json to load the data from the file into a dictionary;
- Return the dictionary;

<!-- grade_id: cell-ce9179c0612fd6b0, points: 0, grade: false, solution: true -->
```python
def import_data():
    with open('data/countries.json', 'r') as f:
        data = json.load(f)
    return data

```

<!-- grade_id: cell-78d54c806fb28266, points: 2, grade: true, solution: false -->
```python
countries = import_data()
assert isinstance(countries, dict), "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
assert countries["Albania"]["Continent"] == "Europe", "Not correct. Keep trying."
assert countries["Vietnam"]["Area"] == 329560, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries_backup = import_data()
```

Output:
```
---- All asserts passed ----
```

## Question 2

The dictionary that you built contains information about a serie of countries. So from now on, we’ll call it the `countries` dictionary.
In the `countries` dictionary, each country has a field called `Population` (be aware that this is case sensitive) that gives the number of people living in that country. Since this is a count, the data should have integer type. However, we have found some numbers with type of *float*. Let's try to fix this!

Complete the function `convert_population_to_int`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Population` field from float to int type for all the countries, using a `for` loop;
- Return the updated `countries` dictionary;

<!-- grade_id: cell-28d89deeef261406, points: 0, grade: false, solution: true -->
```python
def convert_population_to_int(countries):
   
    
    for key in countries:
        if not isinstance (countries[key]["Population"],int):
            countries[key]["Population"] = int(countries[key]["Population"])
    return countries

convert_population_to_int(import_data())

```

Output:
```
{'Afghanistan': {'Continent': 'Asia',
  'Population': 31056997,
  'Area': 647500,
  'GDP': 700.0,
  'Literacy': 36.0,
  'Phones': 3.22},
 'Albania': {'Continent': 'Europe',
  'Population': 3581655,
  'Area': 28748,
  'GDP': 4500.0,
  'Literacy': 86.5,
  'Phones': 71.2},
 'Algeria': {'Continent': 'Africa',
  'Population': 32930091,
  'Area': 2381740,
  'GDP': 6000.0,
  'Literacy': 70.0,
  'Phones': 78.1},
 'American Samoa': {'Continent': 'Oceania',
  'Population': 57794,
  'Area': 199,
  'GDP': 8000.0,
  'Literacy': 97.0,
  'Phones': 259.54},
 'Andorra': {'Continent': 'Europe',
  'Population': 71201,
  'Area': 468,
  'GDP': 19000.0,
  'Literacy': 100.0,
  'Phones': 497.18},
 'Angola': {'Continent': 'Africa',
  'Population': 12127071,
  'Area': 1246700,
  'GDP': 1900.0,
  'Literacy': 42.0,
  'Phones': 7.78},
 'Anguilla': {'Continent': 'America',
  'Population': 13477,
  'Area': 102,
  'GDP': 8600.0,
  'Literacy': 95.0,
  'Phones': 460.04},
 'Antigua & Barbuda': {'Continent': 'America',
 
```

<!-- grade_id: cell-33b501c79f675ab0, points: 2, grade: true, solution: false -->
```python
countries = convert_population_to_int(countries)
assert len(countries) == 194, "Not correct. Keep trying."
source = inspect.getsource(convert_population_to_int)
assert "for" in source, "Not correct. Keep trying."
assert _hash(str(countries)) == "4e2563f26e", "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 3

In the countries dictionary, each country has a field called `Area`. This is the area of the country given in square miles.

Complete the function `convert_area_to_sq_km`. This function should:
- Receive the `countries` dictionary as argument;
- Convert the `Area` field from square miles to square kilometres;
 - Use the following conversion: 1 sq mi = 2.58999 sq km;
 - Round the result to 1 decimal digit;
- Return the updated `countries` dictionary;

<!-- grade_id: cell-3fd27539b41bed5c, points: 0, grade: false, solution: true -->
```python
def convert_area_to_sq_km(countries):
    for key in countries:
        countries[key]["Area"] = round(((countries[key]["Area"]) * 2.58999),1)
    
    print (countries["Haiti"]["Area"])
    return countries
    

convert_area_to_sq_km(import_data())
```

Output:
```
71872.2
{'Afghanistan': {'Continent': 'Asia',
  'Population': 31056997.0,
  'Area': 1677018.5,
  'GDP': 700.0,
  'Literacy': 36.0,
  'Phones': 3.22},
 'Albania': {'Continent': 'Europe',
  'Population': 3581655.0,
  'Area': 74457.0,
  'GDP': 4500.0,
  'Literacy': 86.5,
  'Phones': 71.2},
 'Algeria': {'Continent': 'Africa',
  'Population': 32930091.0,
  'Area': 6168682.8,
  'GDP': 6000.0,
  'Literacy': 70.0,
  'Phones': 78.1},
 'American Samoa': {'Continent': 'Oceania',
  'Population': 57794.0,
  'Area': 515.4,
  'GDP': 8000.0,
  'Literacy': 97.0,
  'Phones': 259.54},
 'Andorra': {'Continent': 'Europe',
  'Population': 71201.0,
  'Area': 1212.1,
  'GDP': 19000.0,
  'Literacy': 100.0,
  'Phones': 497.18},
 'Angola': {'Continent': 'Africa',
  'Population': 12127071.0,
  'Area': 3228940.5,
  'GDP': 1900.0,
  'Literacy': 42.0,
  'Phones': 7.78},
 'Anguilla': {'Continent': 'America',
  'Population': 13477.0,
  'Area': 264.2,
  'GDP': 8600.0,
  'Literacy': 95.0,
  'Phones': 460.04},
 'Antigua 
```

<!-- grade_id: cell-b77708fde98d925a, points: 2, grade: true, solution: false -->
```python
countries = convert_area_to_sq_km(countries)
for country in countries.keys():
    assert isinstance(countries[country]["Area"], float), \
            "Not correct. Keep trying."
assert len(countries) == 194, "Not correct. Keep trying."
assert math.isclose(countries["Haiti"]["Area"], 71872.2, abs_tol=0.1), \
       "Not correct. Keep trying."
assert math.isclose(countries["Morocco"]["Area"], 1156560.0, abs_tol=0.1), \
       "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
71872.2
---- All asserts passed ----
```

## Question 4

Complete the function `get_europe_countries`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with the names of the countries in Europe sorted by alphabetical order (from A to Z);
 - Use the field called `Continent`;
 - Use list comprehension;
 - Use the sort method from the list data type;
- Return the list;

<!-- grade_id: cell-ee1f38ef730c2004, points: 0, grade: false, solution: true -->
```python
def get_europe_countries(countries):
    lista = [key for key in countries if countries[key]["Continent"] == "Europe"]
    lista.sort()
    print(type(lista))
    return lista


    
get_europe_countries(import_data())
```

Output:
```
<class 'list'>
['Albania',
 'Andorra',
 'Austria',
 'Belgium',
 'Bulgaria',
 'Croatia',
 'Czech Republic',
 'Denmark',
 'Estonia',
 'Finland',
 'France',
 'Germany',
 'Greece',
 'Hungary',
 'Iceland',
 'Ireland',
 'Italy',
 'Latvia',
 'Liechtenstein',
 'Lithuania',
 'Luxembourg',
 'Malta',
 'Monaco',
 'Netherlands',
 'Norway',
 'Poland',
 'Portugal',
 'Romania',
 'San Marino',
 'Serbia',
 'Slovenia',
 'Spain',
 'Sweden',
 'Switzerland',
 'United Kingdom']
```

<!-- grade_id: cell-48bf94a05b056aa0, points: 2, grade: true, solution: false -->
```python
countries_list = get_europe_countries(countries)
assert isinstance(countries_list, list), "Not correct. Keep trying."
assert len(countries_list) == 35, "Not correct. Keep trying."
assert countries_list[0] == "Albania", "Not correct. Keep trying."
assert countries_list[5] == "Croatia", "Not correct. Keep trying."
assert countries_list[34] == "United Kingdom", "Not correct. Keep trying."
source = inspect.getsource(get_europe_countries)
assert "for" in source, "Not correct. Keep trying."
assert "[" in source, "Not correct. Keep trying."
assert "in" in source, "Not correct. Keep trying."
assert ".sort(" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
<class 'list'>
---- All asserts passed ----
```

## Question 5

Complete the function `get_literacy_levels_by_continent`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a continent name as argument;
- For the countries in the continent (received as argument), compute the literacy level as:
 - literacy in [0, 25[ % - VERY_LOW;
 - literacy in [25, 50[ % - LOW;
 - literacy in [50, 70[ % - MEDIUM;
 - literacy in [70, 90[ % - HIGH;
 - literacy in [90, 100] % - VERY_HIGH;
- Return a list of tuples like: [(country_1, literacy_1, literacy_level_1), ..., (country_n, literacy_n, literacy_level_n), ...];

<!-- grade_id: cell-3911f29ed53cd968, points: 0, grade: false, solution: true -->
```python
def get_literacy_levels_by_continent(countries, continent):
    literacy_list = []
    
    for key in countries:
        if countries[key]["Continent"] == continent:
            if countries[key]["Literacy"] < 25:
                literacy_list.append((key,countries[key]["Literacy"],"VERY_LOW"))
            elif countries[key]["Literacy"] < 50:
                literacy_list.append((key,countries[key]["Literacy"],"LOW"))
            elif countries[key]["Literacy"] < 70:
                literacy_list.append((key,countries[key]["Literacy"],"MEDIUM"))
            elif countries[key]["Literacy"] < 90:
                literacy_list.append((key,countries[key]["Literacy"],"HIGH"))
            else:
                literacy_list.append((key,countries[key]["Literacy"],"VERY_HIGH"))

    return literacy_list
            
    
    
    
    
get_literacy_levels_by_continent(import_data(),"Africa")
```

Output:
```
[('Algeria', 70.0, 'HIGH'),
 ('Angola', 42.0, 'LOW'),
 ('Benin', 40.9, 'LOW'),
 ('Botswana', 79.8, 'HIGH'),
 ('Burkina Faso', 26.6, 'LOW'),
 ('Burundi', 51.6, 'MEDIUM'),
 ('Cameroon', 79.0, 'HIGH'),
 ('Cape Verde', 76.6, 'HIGH'),
 ('Central African Rep.', 51.0, 'MEDIUM'),
 ('Chad', 47.5, 'LOW'),
 ('Comoros', 56.5, 'MEDIUM'),
 ('Congo, Dem. Rep.', 65.5, 'MEDIUM'),
 ('Congo, Repub. of the', 83.8, 'HIGH'),
 ("Cote d'Ivoire", 50.9, 'MEDIUM'),
 ('Djibouti', 67.9, 'MEDIUM'),
 ('Egypt', 57.7, 'MEDIUM'),
 ('Equatorial Guinea', 85.7, 'HIGH'),
 ('Eritrea', 58.6, 'MEDIUM'),
 ('Ethiopia', 42.7, 'LOW'),
 ('Gabon', 63.2, 'MEDIUM'),
 ('Gambia, The', 40.1, 'LOW'),
 ('Ghana', 74.8, 'HIGH'),
 ('Guinea', 35.9, 'LOW'),
 ('Guinea-Bissau', 42.4, 'LOW'),
 ('Kenya', 85.1, 'HIGH'),
 ('Lesotho', 84.8, 'HIGH'),
 ('Liberia', 57.5, 'MEDIUM'),
 ('Libya', 82.6, 'HIGH'),
 ('Madagascar', 68.9, 'MEDIUM'),
 ('Malawi', 62.7, 'MEDIUM'),
 ('Mali', 46.4, 'LOW'),
 ('Mauritania', 41.7, 'LOW'),
 ('Mauritius', 85.6, 'HIGH'),
 (
```

<!-- grade_id: cell-fa41a51f9ff00d7c, points: 2, grade: true, solution: false -->
```python
for continent in ["Europe","Africa"]:
    literacy_tuple = get_literacy_levels_by_continent(countries, continent)    
    for literacy in literacy_tuple:
        assert len(literacy) == 3, "Not correct. Keep trying."
        assert isinstance(literacy[0], str), "Not correct. Keep trying."
        assert isinstance(literacy[1], float), "Not correct. Keep trying."
        assert isinstance(literacy[2], str), "Not correct. Keep trying."

    if continent == "Europe":
        assert len(literacy_tuple) == 35, "Not correct. Keep trying."
        assert literacy_tuple[21][0] == 'Malta', "Not correct. Keep trying."
        assert math.isclose(literacy_tuple[21][1], 92.8, abs_tol=0.1), \
               "Not correct. Keep trying."
        assert literacy_tuple[21][2] == 'VERY_HIGH', "Not correct. Keep trying."
        assert literacy_tuple[1][2] == 'VERY_HIGH', "Not correct. Keep trying."
        assert literacy_tuple[0][2] == 'HIGH', "Not correct. Keep trying."
        assert literacy_tuple[-1][2] == 'VERY_HIGH', "Not correct. Keep trying."

    if continent == "Africa":
        assert len(literacy_tuple) == 55, "Not correct. Keep trying."
        assert literacy_tuple[14][0] == 'Djibouti', "Not correct. Keep trying."
        assert math.isclose(literacy_tuple[14][1], 67.9, abs_tol=0.1), \
               "Not correct. Keep trying."
        assert literacy_tuple[14][2] == 'MEDIUM', "Not correct. Keep trying."
        assert literacy_tuple[1][2] == 'LOW', "Not correct. Keep trying."
        assert literacy_tuple[0][2] == 'HIGH', "Not correct. Keep trying."
        assert literacy_tuple[-1][2] == 'VERY_HIGH', "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 6

Complete the function `get_country_codes`. This function should:
- Receive the `countries` dictionary as argument;
- Build a list with all the countries names;
- Using `map` and a `lambda` function, convert the country names in country codes by selecting the 3 first letters of the country;
- Returns this list;

Note: Don't worry about duplicates on the result.

<!-- grade_id: cell-bd068242631b4909, points: 0, grade: false, solution: true -->
```python
def get_country_codes(countries):
    country_names = list(countries.keys())
    country_codes = list(map(lambda name: name[:3], country_names))
    return country_codes
    
    
    
get_country_codes(import_data())
```

Output:
```
['Afg',
 'Alb',
 'Alg',
 'Ame',
 'And',
 'Ang',
 'Ang',
 'Ant',
 'Arg',
 'Aru',
 'Aus',
 'Aus',
 'Bah',
 'Bah',
 'Ban',
 'Bar',
 'Bel',
 'Bel',
 'Ben',
 'Ber',
 'Bhu',
 'Bol',
 'Bot',
 'Bra',
 'Bri',
 'Bru',
 'Bul',
 'Bur',
 'Bur',
 'Bur',
 'Cam',
 'Cam',
 'Can',
 'Cap',
 'Cay',
 'Cen',
 'Cha',
 'Chi',
 'Chi',
 'Col',
 'Com',
 'Con',
 'Con',
 'Coo',
 'Cos',
 'Cot',
 'Cro',
 'Cub',
 'Cze',
 'Den',
 'Dji',
 'Dom',
 'Dom',
 'Ecu',
 'Egy',
 'El ',
 'Equ',
 'Eri',
 'Est',
 'Eth',
 'Fij',
 'Fin',
 'Fra',
 'Fre',
 'Fre',
 'Gab',
 'Gam',
 'Ger',
 'Gha',
 'Gre',
 'Gre',
 'Gua',
 'Gua',
 'Gua',
 'Gui',
 'Gui',
 'Guy',
 'Hai',
 'Hon',
 'Hon',
 'Hun',
 'Ice',
 'Ind',
 'Ind',
 'Ira',
 'Ira',
 'Ire',
 'Isr',
 'Ita',
 'Jam',
 'Jap',
 'Jor',
 'Ken',
 'Kor',
 'Kor',
 'Kuw',
 'Lao',
 'Lat',
 'Leb',
 'Les',
 'Lib',
 'Lib',
 'Lie',
 'Lit',
 'Lux',
 'Mac',
 'Mad',
 'Mal',
 'Mal',
 'Mal',
 'Mal',
 'Mal',
 'Mar',
 'Mar',
 'Mau',
 'Mau',
 'Mex',
 'Mic',
 'Mon',
 'Mon',
 'Mor',
 'Moz',
 'Nam',
 'Nep',
 'Net',

```

<!-- grade_id: cell-5f1eaf642616c8a1, points: 2, grade: true, solution: false -->
```python
country_codes = get_country_codes(countries) 
assert len(country_codes) == 194, "Not correct. Keep trying."
for code in country_codes:
    assert isinstance(code, str), "Not correct. Keep trying."
    assert len(code) == 3, "Not correct. Keep trying."
assert country_codes[134] == "Pak", "Not correct. Keep trying."
assert country_codes[62] == "Fra", "Not correct. Keep trying."

source = inspect.getsource(get_country_codes)
assert "map" in source, "Not correct. Keep trying."
assert "lambda" in source, "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 7

The `Country` class is going to be used to represent countries and to return information about them.

### Question 7.1

Complete the `get_population_in_millions` method. This method should:
- Return the population of the country that the class represents;
- The population should be returned in millions with 2 decimal digits;

<!-- grade_id: cell-9b9ae4ba8e61e508, points: 0, grade: false, solution: true -->
```python
class Country:
    def __init__(self, country_name, population, continent):
        self.country_name = country_name
        self.population = population
        self.continent = continent

    def get_population_in_millions(self):
        return round(self.population/1000000,2)
        
```

<!-- grade_id: cell-9e28e403f317f23d, points: 2, grade: true, solution: false -->
```python
country = Country("Penguinea", 12345678, "Antarctica")
assert isinstance(country, Country), "Not correct. Keep trying."
assert isinstance(country.get_population_in_millions(), float), "Not correct. Keep trying."
assert math.isclose(country.get_population_in_millions(), 12.35, abs_tol=0.01),"Not correct. Keep trying."

country = Country("Walrussia", 8344567, "Antarctica") 
assert math.isclose(country.get_population_in_millions(), 8.34, abs_tol=0.01), "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

## Question 7.2

Complete the `get_country_population`. This function should:
- Receive the `countries` dictionary as argument;
- Receive a country name as argument;
- Try to find the country (received as argument) in the countries dictionary;
- If the country doesn’t exist in the countries dictionary, catch the KeyError, print a statement with the information that there is no information for that country and return `None`;
- If the country exists, create an object of the `Country` class;
- Call the `get_population_in_millions` method for the country;
- Return the population in millions value;

<!-- grade_id: cell-803b263dbee2e300, points: 0, grade: false, solution: true -->
```python
def get_country_population(countries, country_name):
    if country_name in list(countries.keys()):
        test = Country(country_name,countries[country_name]["Population"],countries[country_name]["Continent"])
        return test.get_population_in_millions()
    else:
        return None
    
    
    
get_country_population(import_data(),"Turkey")
```

Output:
```
70.41
```

<!-- grade_id: cell-0d0db73f5256c8e5, points: 6, grade: true, solution: false -->
```python
assert math.isclose(get_country_population(countries, "Turkey"), 70.41, abs_tol=0.01), "Not correct. Keep trying."
assert get_country_population(countries, "Narnia") is None,  "Not correct. Keep trying."
print("---- All asserts passed ---- ")
countries = countries_backup
```

Output:
```
---- All asserts passed ----
```

### Submit your work!

To submit your work, [follow the steps here, in the step \"Grading the Exercise Notebook\"!](https://github.com/LDSSA/ds-prep-course-2023/blob/main/weekly-workflow.md#link-to-grading)